./isabelle-mirabelle/BO_cvc42/x2020_08_03_15_24_49_644_6536112.cvc5oldpf:(step t1 (cl (= (= true false) false)) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_16_50_48_855_5281470.cvc5oldpf:(step t4 (cl (= (= (arg_d ua_d) sa_d) (= sa_d (arg_d ua_d)))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_16_51_48_838_5354134.cvc5oldpf:(step t4 (cl (= (= (arg_d ua_d) sa_d) (= sa_d (arg_d ua_d)))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_16_51_55_014_5369464.cvc5oldpf:(step t1 (cl (= (= true false) false)) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_15_56_39_113_7636710.cvc5oldpf:(step t1 (cl (= (= t2_d sa_d) (= sa_d t2_d))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_15_56_39_113_7636710.cvc5oldpf:(step t10 (cl (= (= (arg_d ta_d) sa_d) (= sa_d (arg_d ta_d)))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_15_23_28_446_6397402.cvc5oldpf:(step t4 (cl (= (= (arg_d ua_d) sa_d) (= sa_d (arg_d ua_d)))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_16_51_50_546_5359382.cvc5oldpf:(step t1 (cl (= (= false true) false)) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_16_51_50_546_5359382.cvc5oldpf:(step t15 (cl (= (not (not (comp_hd_d gt_sym_d ground_heads_var_d (head_d ua_d) (head_d sa_d)))) (comp_hd_d gt_sym_d ground_heads_var_d (head_d ua_d) (head_d sa_d)))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_16_42_01_808_4961238.cvc5oldpf:(step t4 (cl (= (= (arg_d ua_d) sa_d) (= sa_d (arg_d ua_d)))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_16_40_35_239_4899098.cvc5oldpf:(step t1 (cl (= (not (not (= (wt_d ground_heads_var_d delta_d arity_sym_d wt_sym_d ua_d) (wt_d ground_heads_var_d delta_d arity_sym_d wt_sym_d sa_d)))) (= (wt_d ground_heads_var_d delta_d arity_sym_d wt_sym_d ua_d) (wt_d ground_heads_var_d delta_d arity_sym_d wt_sym_d sa_d)))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_15_14_36_388_6025794.cvc5oldpf:(step t4 (cl (= (= (arg_d ua_d) sa_d) (= sa_d (arg_d ua_d)))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_15_24_43_355_6522276.cvc5oldpf:(step t1 (cl (= (= false true) false)) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_15_24_43_355_6522276.cvc5oldpf:(step t15 (cl (= (not (not (comp_hd_d gt_sym_d ground_heads_var_d (head_d ua_d) (head_d sa_d)))) (comp_hd_d gt_sym_d ground_heads_var_d (head_d ua_d) (head_d sa_d)))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_15_06_02_107_5638220.cvc5oldpf:(step t1 (cl (= (= true false) false)) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_16_00_26_055_7846716.cvc5oldpf:(step t1 (cl (= (not (not (= (head_d ta_d) (head_d sa_d)))) (= (head_d ta_d) (head_d sa_d)))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_15_22_36_035_6337544.cvc5oldpf:(step t4 (cl (= (= (arg_d ua_d) sa_d) (= sa_d (arg_d ua_d)))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_16_51_08_268_5321078.cvc5oldpf:(step t1 (cl (= (= false true) false)) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_15_23_49_010_6452792.cvc5oldpf:(step t1 (cl (= (= false true) false)) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_17_12_58_132_6181512.cvc5oldpf:(step t1 (cl (= (= false true) false)) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_17_08_43_615_6004000.cvc5oldpf:(step t1 (cl (= (= t2_d sa_d) (= sa_d t2_d))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_17_08_43_615_6004000.cvc5oldpf:(step t10 (cl (= (= (arg_d ta_d) sa_d) (= sa_d (arg_d ta_d)))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_15_24_41_942_6515590.cvc5oldpf:(step t4 (cl (= (= (arg_d ua_d) sa_d) (= sa_d (arg_d ua_d)))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_16_50_04_968_5248672.cvc5oldpf:(step t4 (cl (= (= (arg_d ua_d) sa_d) (= sa_d (arg_d ua_d)))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_17_15_27_694_6279318.cvc5oldpf:(step t1 (cl (= (= (vars_d t_d) bot_d) (= bot_d (vars_d t_d)))) :rule all_simplify)
./isabelle-mirabelle/BO_cvc42/x2020_08_03_17_15_27_694_6279318.cvc5oldpf:(step t3 (cl (= (= (vars_d s_d) bot_d) (= bot_d (vars_d s_d)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_09_37_53_950_9571100.cvc5oldpf:(step t1 (cl (= (not (not (= c_d zero_d))) (= c_d zero_d))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_17_45_18_358_5383148.cvc5oldpf:(step t1 (cl (= (not (not (member_d bot_d a_d))) (member_d bot_d a_d))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_32_10_259_5099720.cvc5oldpf:(step t1 (cl (= (< (+ (+ e1_d (* -1 e2_d)) (* -1 (+ e1_d (* -1 e2_d)))) (+ 0 (* -1 1))) (not (>= (+ (+ e1_d (* -1 e2_d)) (* -1 (+ e1_d (* -1 e2_d)))) (+ 0 (* -1 1)))))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_32_10_259_5099720.cvc5oldpf:(step t2 (cl (= (+ (+ e1_d (* -1 e2_d)) (* -1 (+ e1_d (* -1 e2_d)))) (+ e1_d (* -1 e2_d) (* -1 (+ e1_d (* -1 e2_d)))))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_32_10_259_5099720.cvc5oldpf:(step t5 (cl (= (* -1 (+ e1_d (* -1 e2_d))) (+ (* -1 e1_d) e2_d))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_32_10_259_5099720.cvc5oldpf:(step t7 (cl (= (+ e1_d (* -1 e2_d) (+ (* -1 e1_d) e2_d)) 0)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_32_10_259_5099720.cvc5oldpf:(step t11 (cl (= (* -1 1) -1)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_32_10_259_5099720.cvc5oldpf:(step t13 (cl (= (+ 0 -1) -1)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_32_10_259_5099720.cvc5oldpf:(step t16 (cl (= (>= 0 -1) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_32_10_259_5099720.cvc5oldpf:(step t19 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_32_10_259_5099720.cvc5oldpf:(step t24 (cl (= (< (+ e1_d (* -1 e2_d)) 0) (not (>= (+ e1_d (* -1 e2_d)) 0)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_32_10_259_5099720.cvc5oldpf:(step t27 (cl (= (<= e2_d e1_d) (>= (+ e1_d (* -1 e2_d)) 0))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_32_10_259_5099720.cvc5oldpf:(step t32 (cl (= (= thesis_d false) (not thesis_d))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_32_10_259_5099720.cvc5oldpf:(step t36 (cl (= (=> (>= (+ e1_d (* -1 e2_d)) 0) false) (not (>= (+ e1_d (* -1 e2_d)) 0)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_32_10_259_5099720.cvc5oldpf:(step t49 (cl (= (<= e1_d e2_d) (not (>= (+ e1_d (* -1 e2_d)) 1)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_32_10_259_5099720.cvc5oldpf:(step t53 (cl (= (=> (not (>= (+ e1_d (* -1 e2_d)) 1)) false) (>= (+ e1_d (* -1 e2_d)) 1))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_16_51_570_4983100.cvc5oldpf:(step t15 (cl (= (not (not (= a_d b_d))) (= a_d b_d))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_16_51_570_4983100.cvc5oldpf:(step t28 (cl (= (not (not b_d)) b_d)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_16_51_570_4983100.cvc5oldpf:(step t50 (cl (= (not (not a_d)) a_d)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_09_33_39_249_9488436.cvc5oldpf:(step t1 (cl (= (not (not (= c_d zero_d))) (= c_d zero_d))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_09_32_04_209_9474402.cvc5oldpf:(step t1 (cl (= (not (not (= c_d zero_d))) (= c_d zero_d))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_23_22_32_196_4983000.cvc5oldpf:(step t6 (cl (= (= (f_d n_d) (g_d n_d)) (= (g_d n_d) (f_d n_d)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_07_48_27_625_8447402.cvc5oldpf:(step t1 (cl (= (= (fold_d max_d (tl_d (keys_d ta_d)) (hd_d (keys_d ta_d))) (fold_d max_d (tl_d (keys_d ta_d)) (hd_d (keys_d ta_d)))) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_07_48_27_625_8447402.cvc5oldpf:(step t3 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_05_38_13_554_5028546.cvc5oldpf:(step t1 (cl (= (= (pair_d bot_d bot_da) (pair_d bot_d bot_da)) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_05_38_13_554_5028546.cvc5oldpf:(step t3 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_05_38_10_024_5021652.cvc5oldpf:(step t1 (cl (= (= (pair_d top_d top_da) (pair_d top_d top_da)) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_05_38_10_024_5021652.cvc5oldpf:(step t3 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_35_20_861_5083584.cvc5oldpf:(step t1 (cl (= (< (+ n_d (* -1 n_d)) (+ 1 (* -1 2))) (not (>= (+ n_d (* -1 n_d)) (+ 1 (* -1 2)))))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_35_20_861_5083584.cvc5oldpf:(step t2 (cl (= (+ n_d (* -1 n_d)) 0)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_35_20_861_5083584.cvc5oldpf:(step t4 (cl (= (* -1 2) -2)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_35_20_861_5083584.cvc5oldpf:(step t6 (cl (= (+ 1 -2) -1)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_35_20_861_5083584.cvc5oldpf:(step t9 (cl (= (>= 0 -1) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_35_20_861_5083584.cvc5oldpf:(step t12 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_35_20_861_5083584.cvc5oldpf:(step t17 (cl (= (< n_d 1) (not (>= n_d 1)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_35_20_861_5083584.cvc5oldpf:(step t20 (cl (= (< 0 n_d) (not (>= 0 n_d)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_35_20_861_5083584.cvc5oldpf:(step t21 (cl (= (>= 0 n_d) (not (>= n_d 1)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_35_20_861_5083584.cvc5oldpf:(step t23 (cl (= (not (not (>= n_d 1))) (>= n_d 1))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_35_20_861_5083584.cvc5oldpf:(step t37 (cl (= (< 1 n_d) (not (>= 1 n_d)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_35_20_861_5083584.cvc5oldpf:(step t38 (cl (= (>= 1 n_d) (not (>= n_d 2)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_23_15_35_20_861_5083584.cvc5oldpf:(step t40 (cl (= (not (not (>= n_d 2))) (>= n_d 2))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_50_06_230_5063492.cvc5oldpf:(step t2 (cl (= (= (image_d f_d s_d) (image_d f_d s_d)) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_00_50_06_230_5063492.cvc5oldpf:(step t4 (cl (= (and (inj_on_d f_d s_d) true) (inj_on_d f_d s_d))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_07_41_03_832_8304788.cvc5oldpf:(step t1 (cl (= (= empty_d empty_d) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_07_41_03_832_8304788.cvc5oldpf:(step t3 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_09_08_53_363_5071014.cvc5oldpf:(step t5 (cl (= (= (class_d (pick_d (class_d a_da))) (class_d a_da)) (= (class_d a_da) (class_d (pick_d (class_d a_da)))))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_09_08_53_363_5071014.cvc5oldpf:(step t9 (cl (= (= (class_d (pick_d a_d)) a_d) (= a_d (class_d (pick_d a_d))))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_02_50_45_899_5005750.cvc5oldpf:(step t1 (cl (= (ite p_d true q_d) (or p_d q_d))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_02_50_45_899_5005750.cvc5oldpf:(step t4 (cl (= (= (or p_d q_d) (or p_d q_d)) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_02_50_45_899_5005750.cvc5oldpf:(step t7 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_09_09_13_580_5081634.cvc5oldpf:(step t1 (cl (= (= false true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_07_28_08_034_7411138.cvc5oldpf:(step t1 (cl (= (= empty_d empty_d) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_07_28_08_034_7411138.cvc5oldpf:(step t3 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_07_47_03_336_8415396.cvc5oldpf:(step t1 (cl (= (= (fold_d min_d (tl_d (keys_d ta_d)) (hd_d (keys_d ta_d))) (fold_d min_d (tl_d (keys_d ta_d)) (hd_d (keys_d ta_d)))) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_verit/x2020_07_24_07_47_03_336_8415396.cvc5oldpf:(step t3 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_09_13_25_289_5134680.cvc5oldpf:(step t5 (cl (= (= (class_d (pick_d (class_d a_da))) (class_d a_da)) (= (class_d a_da) (class_d (pick_d (class_d a_da)))))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_09_13_25_289_5134680.cvc5oldpf:(step t9 (cl (= (= (class_d (pick_d a_d)) a_d) (= a_d (class_d (pick_d a_d))))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_01_29_51_707_5106282.cvc5oldpf:(step t2 (cl (= (= (image_d f_d s_d) (image_d f_d s_d)) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_01_29_51_707_5106282.cvc5oldpf:(step t4 (cl (= (and (inj_on_d f_d s_d) true) (inj_on_d f_d s_d))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_08_28_06_559_5486688.cvc5oldpf:(step t1 (cl (= (= empty_d empty_d) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_08_28_06_559_5486688.cvc5oldpf:(step t3 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_15_23_16_143_4983100.cvc5oldpf:(step t15 (cl (= (not (not (= a_d b_d))) (= a_d b_d))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_15_23_16_143_4983100.cvc5oldpf:(step t28 (cl (= (not (not b_d)) b_d)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_15_23_16_143_4983100.cvc5oldpf:(step t50 (cl (= (not (not a_d)) a_d)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_05_49_32_934_5066708.cvc5oldpf:(step t1 (cl (= (= (pair_d bot_d bot_da) (pair_d bot_d bot_da)) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_05_49_32_934_5066708.cvc5oldpf:(step t3 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_08_36_42_549_5864026.cvc5oldpf:(step t1 (cl (= (= (fold_d min_d (tl_d (keys_d ta_d)) (hd_d (keys_d ta_d))) (fold_d min_d (tl_d (keys_d ta_d)) (hd_d (keys_d ta_d)))) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_08_36_42_549_5864026.cvc5oldpf:(step t3 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_18_31_06_963_5320700.cvc5oldpf:(step t1 (cl (= (not (not (member_d bot_d a_d))) (member_d bot_d a_d))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_08_38_08_504_5931826.cvc5oldpf:(step t1 (cl (= (= (fold_d max_d (tl_d (keys_d ta_d)) (hd_d (keys_d ta_d))) (fold_d max_d (tl_d (keys_d ta_d)) (hd_d (keys_d ta_d)))) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_08_38_08_504_5931826.cvc5oldpf:(step t3 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_23_22_38_994_4983000.cvc5oldpf:(step t6 (cl (= (= (f_d n_d) (g_d n_d)) (= (g_d n_d) (f_d n_d)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_16_01_56_200_5114158.cvc5oldpf:(step t1 (cl (= (< (+ n_d (* -1 n_d)) (+ 1 (* -1 2))) (not (>= (+ n_d (* -1 n_d)) (+ 1 (* -1 2)))))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_16_01_56_200_5114158.cvc5oldpf:(step t2 (cl (= (+ n_d (* -1 n_d)) 0)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_16_01_56_200_5114158.cvc5oldpf:(step t4 (cl (= (* -1 2) -2)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_16_01_56_200_5114158.cvc5oldpf:(step t6 (cl (= (+ 1 -2) -1)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_16_01_56_200_5114158.cvc5oldpf:(step t9 (cl (= (>= 0 -1) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_16_01_56_200_5114158.cvc5oldpf:(step t12 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_16_01_56_200_5114158.cvc5oldpf:(step t17 (cl (= (< n_d 1) (not (>= n_d 1)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_16_01_56_200_5114158.cvc5oldpf:(step t20 (cl (= (< 0 n_d) (not (>= 0 n_d)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_16_01_56_200_5114158.cvc5oldpf:(step t21 (cl (= (>= 0 n_d) (not (>= n_d 1)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_16_01_56_200_5114158.cvc5oldpf:(step t23 (cl (= (not (not (>= n_d 1))) (>= n_d 1))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_16_01_56_200_5114158.cvc5oldpf:(step t37 (cl (= (< 1 n_d) (not (>= 1 n_d)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_16_01_56_200_5114158.cvc5oldpf:(step t38 (cl (= (>= 1 n_d) (not (>= n_d 2)))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_23_16_01_56_200_5114158.cvc5oldpf:(step t40 (cl (= (not (not (>= n_d 2))) (>= n_d 2))) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_08_33_14_874_5693428.cvc5oldpf:(step t1 (cl (= (= empty_d empty_d) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_08_33_14_874_5693428.cvc5oldpf:(step t3 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_05_49_28_655_5053232.cvc5oldpf:(step t1 (cl (= (= (pair_d top_d top_da) (pair_d top_d top_da)) true)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_05_49_28_655_5053232.cvc5oldpf:(step t3 (cl (= (not true) false)) :rule all_simplify)
./isabelle-mirabelle/HOL-Library/smt_cvc4/x2020_07_24_09_13_44_183_5157124.cvc5oldpf:(step t1 (cl (= (= false true) false)) :rule all_simplify)
./isabelle-mirabelle/Green_cvc42/x2020_07_31_07_55_13_484_7016530.cvc5oldpf:(step t2 (cl (= (not (not (islimpt_d x_d (minus_d s_d t_d)))) (islimpt_d x_d (minus_d s_d t_d)))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_03_33_34_830_8439854.cvc5oldpf:(step t2 (cl (= (= (s_M_d subst_atm_d s_d m_d (plus_d d_d (image_mset_d neg_d (mset_d as_d)))) zero_d) (= zero_d (s_M_d subst_atm_d s_d m_d (plus_d d_d (image_mset_d neg_d (mset_d as_d))))))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_02_17_01_007_5595370.cvc5oldpf:(step t1 (cl (= (= (size_da as_d) n_d) (= n_d (size_da as_d)))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_00_24_03_784_9778214.cvc5oldpf:(step t1 (cl (= (not (not (member_d e_d (n_of_state_d (lnth_d sts_d j_d))))) (member_d e_d (n_of_state_d (lnth_d sts_d j_d))))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_04_21_25_416_8839786.cvc5oldpf:(step t1 (cl (= (= (size_d cAs0_d) n_d) (= n_d (size_d cAs0_d)))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_04_21_25_416_8839786.cvc5oldpf:(step t5 (cl (= (= (size_d cAs_d) (size_d cAs0_d)) (= (size_d cAs0_d) (size_d cAs_d)))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_03_48_18_914_8656120.cvc5oldpf:(step t1 (cl (= (not (not (= (nth_d (subst_cls_list_d subst_atm_d cs0_d eta_d) i_d) (nth_d cs_d i_d)))) (= (nth_d (subst_cls_list_d subst_atm_d cs0_d eta_d) i_d) (nth_d cs_d i_d)))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_28_22_15_13_395_5646634.cvc5oldpf:(step t1 (cl (= (= (subst_cls_d subst_atm_d d_d sigma_d) c_d) (= c_d (subst_cls_d subst_atm_d d_d sigma_d)))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_28_21_17_39_123_7537830.cvc5oldpf:(step t1 (cl (= (not (not (and (generalizes_d subst_atm_d c_d c_d) (not (generalizes_d subst_atm_d c_d c_d))))) (and (generalizes_d subst_atm_d c_d c_d) (not (generalizes_d subst_atm_d c_d c_d))))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_04_46_15_427_5720368.cvc5oldpf:(step t1 (cl (= (= true false) false)) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_28_22_03_52_352_5382930.cvc5oldpf:(step t2 (cl (= (not (not (= c_d zero_d))) (= c_d zero_d))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_28_22_02_04_807_5590668.cvc5oldpf:(step t1 (cl (= (not (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))) (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_28_22_02_04_807_5590668.cvc5oldpf:(step t4 (cl (= (=> (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d) false) (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d)))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_03_39_24_709_8527638.cvc5oldpf:(step t1.t2 (cl (= (not (not (= zero_d (s_M_d subst_atm_d s_d m_d (plus_d d_d (image_mset_d neg_d (mset_d as_d))))))) (= zero_d (s_M_d subst_atm_d s_d m_d (plus_d d_d (image_mset_d neg_d (mset_d as_d))))))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_03_39_24_709_8527638.cvc5oldpf:(step t1.t6 (cl (= (= (s_M_d subst_atm_d s_d m_d (plus_d d_d (image_mset_d neg_d (mset_d as_d)))) zero_d) (= zero_d (s_M_d subst_atm_d s_d m_d (plus_d d_d (image_mset_d neg_d (mset_d as_d))))))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_03_39_24_709_8527638.cvc5oldpf:(step t9 (cl (= (not (not (= zero_d (s_M_d subst_atm_d s_d m_d (plus_d d_d (image_mset_d neg_d (mset_d as_d))))))) (= zero_d (s_M_d subst_atm_d s_d m_d (plus_d d_d (image_mset_d neg_d (mset_d as_d))))))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_03_39_24_709_8527638.cvc5oldpf:(step t13 (cl (= (= (s_M_d subst_atm_d s_d m_d (plus_d d_d (image_mset_d neg_d (mset_d as_d)))) zero_d) (= zero_d (s_M_d subst_atm_d s_d m_d (plus_d d_d (image_mset_d neg_d (mset_d as_d))))))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_04_21_27_449_8841898.cvc5oldpf:(step t1 (cl (= (= (size_d eta_s0_d) n_d) (= n_d (size_d eta_s0_d)))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_04_21_27_449_8841898.cvc5oldpf:(step t5 (cl (= (= (size_da cAs_d) (size_d eta_s0_d)) (= (size_d eta_s0_d) (size_da cAs_d)))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_04_23_25_559_8874768.cvc5oldpf:(step t1 (cl (= (= (size_da asa_d) n_d) (= n_d (size_da asa_d)))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_04_00_10_997_8754562.cvc5oldpf:(step t2 (cl (= (not (not (= zero_d (s_M_d subst_atm_d s_d m_d (plus_d d_d (image_mset_d neg_d (mset_d as_d))))))) (= zero_d (s_M_d subst_atm_d s_d m_d (plus_d d_d (image_mset_d neg_d (mset_d as_d))))))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_04_00_10_997_8754562.cvc5oldpf:(step t6 (cl (= (= (s_M_d subst_atm_d s_d m_d (plus_d d_d (image_mset_d neg_d (mset_d as_d)))) zero_d) (= zero_d (s_M_d subst_atm_d s_d m_d (plus_d d_d (image_mset_d neg_d (mset_d as_d))))))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_04_44_56_253_5688612.cvc5oldpf:(step t1 (cl (= (not (not (true_cls_d (interp_d s_d n_d) (plus_d (sum_mset_d (mset_d cs_d)) d_d)))) (true_cls_d (interp_d s_d n_d) (plus_d (sum_mset_d (mset_d cs_d)) d_d)))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_04_36_57_199_5464880.cvc5oldpf:(step t2 (cl (= (not (not (= da_d zero_d))) (= da_d zero_d))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_28_21_17_43_975_7545750.cvc5oldpf:(step t1 (cl (= (not (not (and (subsumes_d subst_atm_d d_d c_d) (not (subsumes_d subst_atm_d c_d d_d))))) (and (subsumes_d subst_atm_d d_d c_d) (not (subsumes_d subst_atm_d c_d d_d))))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_02_29_35_898_5741632.cvc5oldpf:(step t18 (cl (= (= (subst_cls_d subst_atm_d (plus_d (sum_mset_d (mset_d cs_da)) dAa_d) sigma_dc) (subst_cls_d subst_atm_d (plus_d (sum_mset_d (mset_d cs_da)) dAa_d) sigma_dc)) true)) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_02_29_35_898_5741632.cvc5oldpf:(step t20 (cl (= (not true) false)) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_28_21_17_42_834_7543784.cvc5oldpf:(step t1 (cl (= (not (not (and (subsumes_d subst_atm_d c_d c_d) (not (subsumes_d subst_atm_d c_d c_d))))) (and (subsumes_d subst_atm_d c_d c_d) (not (subsumes_d subst_atm_d c_d c_d))))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_03_31_29_566_8412284.cvc5oldpf:(step t1 (cl (= (not (not (= (image_mset_d neg_d (mset_d as_d)) (subst_cls_d subst_atm_d (fun_app_d s_d dA0_d) eta_0_d)))) (= (image_mset_d neg_d (mset_d as_d)) (subst_cls_d subst_atm_d (fun_app_d s_d dA0_d) eta_0_d)))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_28_21_17_40_281_7539822.cvc5oldpf:(step t1 (cl (= (not (not (and (generalizes_d subst_atm_d d_d c_d) (not (generalizes_d subst_atm_d c_d d_d))))) (and (generalizes_d subst_atm_d d_d c_d) (not (generalizes_d subst_atm_d c_d d_d))))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_02_01_16_644_5442968.cvc5oldpf:(step t1 (cl (= (=> (not (true_clss_d (i_of_d n_d) n_d)) false) (true_clss_d (i_of_d n_d) n_d))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_03_46_36_378_8617908.cvc5oldpf:(step t1 (cl (= (= false true) false)) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_28_22_02_03_707_5588814.cvc5oldpf:(step t2 (cl (= (= (s_d d_d) zero_d) (= zero_d (s_d d_d)))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_28_22_02_03_707_5588814.cvc5oldpf:(step t9 (cl (= (=> (less_d (max_d (atms_of_d d_d)) (max_d (atms_of_d d_da))) false) (not (less_d (max_d (atms_of_d d_d)) (max_d (atms_of_d d_da)))))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_28_22_02_03_707_5588814.cvc5oldpf:(step t13 (cl (= (=> (and (member_d d_d n_d) (= (s_d d_d) zero_d)) false) (not (and (member_d d_d n_d) (= (s_d d_d) zero_d))))) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_28_19_42_14_120_5961368.cvc5oldpf:(step t1 (cl (= (= (wf_d (collect_d (case_prod_d r_d))) (wf_d (collect_d (case_prod_d r_d)))) true)) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_28_19_42_14_120_5961368.cvc5oldpf:(step t3 (cl (= (not true) false)) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_03_47_44_675_8643846.cvc5oldpf:(step t7 (cl (= (= (nth_d (subst_cls_lists_d subst_atm_d cs0_d eta_s0_d) i_d) (nth_d (subst_cls_lists_d subst_atm_d cs0_d eta_s0_d) i_d)) true)) :rule all_simplify)
./Ordered_Resolution_Prover_veriT/x2020_07_29_03_47_44_675_8643846.cvc5oldpf:(step t9 (cl (= (not true) false)) :rule all_simplify)
