(assume a0 (not (not (and (generalizes_d subst_atm_d d_d c_d) (not (generalizes_d subst_atm_d c_d d_d))))))
(assume a1 (and (generalizes_d subst_atm_d c_d d_d) (not (generalizes_d subst_atm_d d_d c_d))))
(step t2 (cl (not (not (not (and (generalizes_d subst_atm_d d_d c_d) (not (generalizes_d subst_atm_d c_d d_d)))))) (and (generalizes_d subst_atm_d d_d c_d) (not (generalizes_d subst_atm_d c_d d_d)))) :rule not_not)
(step t3 (cl (and (generalizes_d subst_atm_d d_d c_d) (not (generalizes_d subst_atm_d c_d d_d)))) :rule th_resolution :premises (t2 a0))
(step t5 (cl (not (generalizes_d subst_atm_d c_d d_d))) :rule and :premises (t3))
(step t6 (cl (generalizes_d subst_atm_d c_d d_d)) :rule and :premises (a1))
(step t7 (cl) :rule resolution :premises (t6 t5))
