(assume a0 (not (= (pair_d top_d top_da) (pair_d top_d top_da))))
(step t0 (cl (not (= (not (= (pair_d top_d top_da) (pair_d top_d top_da))) false)) (not (not (= (pair_d top_d top_da) (pair_d top_d top_da)))) false) :rule equiv_pos2)
(step t1 (cl (= (= (pair_d top_d top_da) (pair_d top_d top_da)) true)) :rule all_simplify)
(step t2 (cl (= (not (= (pair_d top_d top_da) (pair_d top_d top_da))) (not true))) :rule cong :premises (t1))
(step t3 (cl (= (not true) false)) :rule all_simplify)
(step t4 (cl (= (not (= (pair_d top_d top_da) (pair_d top_d top_da))) false)) :rule trans :premises (t2 t3))
(step t5 (cl false) :rule resolution :premises (t0 t4 a0))
(step t6 (cl (not false)) :rule false)
(step t7 (cl) :rule resolution :premises (t5 t6))
