(assume a0 (not (less_d zero_d delta_d)))
(assume a1 (= delta_d epsilon_d))
(assume a2 (less_d zero_d epsilon_d))
(step t0 (cl (not (= (not (less_d zero_d delta_d)) (not (less_d zero_d epsilon_d)))) (not (not (less_d zero_d delta_d))) (not (less_d zero_d epsilon_d))) :rule equiv_pos2)
(step t1 (cl (= zero_d zero_d)) :rule refl)
(step t2 (cl (= (less_d zero_d delta_d) (less_d zero_d epsilon_d))) :rule cong :premises (t1 a1))
(step t3 (cl (= (not (less_d zero_d delta_d)) (not (less_d zero_d epsilon_d)))) :rule cong :premises (t2))
(step t4 (cl (not (less_d zero_d epsilon_d))) :rule resolution :premises (t0 t3 a0))
(step t5 (cl) :rule resolution :premises (a2 t4))
