(assume a0 (not thesis_d))
(assume a1 (=> (<= e2_d e1_d) thesis_d))
(assume a2 (=> (<= e1_d e2_d) thesis_d))
(step t4 (cl (not (<= e2_d e1_d)) thesis_d) :rule implies :premises (a1))
(step t5 (cl (not (<= e2_d e1_d))) :rule resolution :premises (t4 a0))
(step t6 (cl (not (<= e1_d e2_d)) thesis_d) :rule implies :premises (a2))
(step t7 (cl (not (<= e1_d e2_d))) :rule resolution :premises (t6 a0))
(step t8 (cl (<= e1_d e2_d) (<= e2_d e1_d)) :rule la_generic :args (1 1))
(step t9 (cl) :rule resolution :premises (t8 t5 t7))
