(assume a0 (not (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))))
(assume a1 (=> (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d) false))
(step t0 (cl (not (= (not (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))) (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))) (not (not (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d)))) (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d)) :rule equiv_pos2)
(step t1 (cl (= (not (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))) (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))) :rule all_simplify)
(step t2 (cl (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d)) :rule resolution :premises (t0 t1 a0))
(step t3 (cl (not (= (=> (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d) false) (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d)))) (not (=> (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d) false)) (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))) :rule equiv_pos2)
(step t4 (cl (= (=> (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d) false) (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d)))) :rule all_simplify)
(step t5 (cl (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))) :rule resolution :premises (t3 t4 a1))
(step t6 (cl) :rule resolution :premises (t2 t5))
