(assume a0 (not (= (wf_d (collect_d (case_prod_d r_d))) (wf_d (collect_d (case_prod_d r_d))))))
(step t0 (cl (not (= (not (= (wf_d (collect_d (case_prod_d r_d))) (wf_d (collect_d (case_prod_d r_d))))) false)) (not (not (= (wf_d (collect_d (case_prod_d r_d))) (wf_d (collect_d (case_prod_d r_d)))))) false) :rule equiv_pos2)
(step t1 (cl (= (= (wf_d (collect_d (case_prod_d r_d))) (wf_d (collect_d (case_prod_d r_d)))) true)) :rule all_simplify)
(step t2 (cl (= (not (= (wf_d (collect_d (case_prod_d r_d))) (wf_d (collect_d (case_prod_d r_d))))) (not true))) :rule cong :premises (t1))
(step t3 (cl (= (not true) false)) :rule all_simplify)
(step t4 (cl (= (not (= (wf_d (collect_d (case_prod_d r_d))) (wf_d (collect_d (case_prod_d r_d))))) false)) :rule trans :premises (t2 t3))
(step t5 (cl false) :rule resolution :premises (t0 t4 a0))
(step t6 (cl (not false)) :rule false)
(step t7 (cl) :rule resolution :premises (t5 t6))
