(assume a0 (not (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))))
(assume a1 (=> (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d) false))
(step t1 (cl (not (= (not (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))) (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))) (not (not (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d)))) (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d)) :rule equiv_pos2)
(step t2 (cl (= (not (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))) (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))) :rule all_simplify :args (not_simplify1 (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d)))
(step t3 (cl (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d)) :rule resolution :premises (t1 t2 a0))
(step t4 (cl (not (= (=> (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d) false) (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d)))) (not (=> (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d) false)) (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))) :rule equiv_pos2)
(step t5 (cl (= (=> (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d) false) (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d)))) :rule all_simplify :args (implies_simplify5 (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d)))
(step t6 (cl (not (true_cls_d (sup_d (image_d (production_d s_d n_d) cc_d)) d_d))) :rule resolution :premises (t4 t5 a1))
(step t7 (cl) :rule resolution :premises (t3 t6))
